## Module 3

#### In this assignment, you will continue working on the movie data from IMDB.
- The data includes movies and ratings from the IMDB website
- Data File(s): imdb.xlsx

#### Data file contains 3 sheets:
- “imdb”: contains records of movies and ratings scraped from IMDB website
- “countries”: contains the country (of origin) names
- “directors”: contains the director names

We have loaded the data as "df" for you. Follow the instructions and finish the rest.

<!-- grade_id: init_test, points: 0, grade: true, solution: false -->
```python
###########################################################
### EXECUTE THIS CELL BEFORE YOU TO TEST YOUR SOLUTIONS ###
###########################################################

import imp, os, sys
sol = imp.load_compiled("solutions", "./solutions.py")
sol.get_solutions("imdb.xlsx")
from nose.tools import assert_equal
from pandas.util.testing import assert_frame_equal, assert_series_equal
```

<!-- grade_id: init_user, points: 0, grade: true, solution: false -->
```python
# Loading the data
import pandas as pd

xls = pd.ExcelFile('imdb.xlsx')
df = xls.parse('imdb')
df_directors = xls.parse('directors')
df_countries = xls.parse('countries')

print("Data Loading Finished.")
```

Output:
```
Data Loading Finished.
```

```python
df.columns
```

Output:
```
Index(['movie_title', 'director_id', 'country_id', 'content_rating',
       'title_year', 'imdb_score', 'gross', 'duration'],
      dtype='object')
```

```python
df_directors.columns
```

Output:
```
Index(['id', 'director_name'], dtype='object')
```

<!-- grade_id: q1, points: 0, grade: false, solution: true -->
```python
""" Q1: 
Join three Dataframes: df, df_directors, and df_countries with an inner join.
Store the joined DataFrames in df.

Here are the steps:
1. Merge df with df_countries and assign it df
2. Merge df with df_directors and assign it to df again
There might be errors if the merge is not in this order, so please be careful.

"""

df = pd.merge(left = pd.merge(left = df, right = df_countries, left_on="country_id", right_on="id"), left_on="director_id", right = df_directors, right_on="id")

df.shape
```

Output:
```
(178, 12)
```

<!-- grade_id: q1_test, points: 2, grade: true, solution: false -->
```python
assert_equal(df.shape, sol.df.shape)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q2, points: 0, grade: false, solution: true -->
```python
""" Q2: 
Save the first ten rows of movie titles in a variable called first10, then print it
"""

first10 = df.movie_title[:10]


```

<!-- grade_id: q2_test, points: 1, grade: true, solution: false -->
```python
assert_series_equal(first10, sol.first10)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q3, points: 0, grade: false, solution: true -->
```python
""" Q3: 
There's an extra character at the end of each movie title. 
Remove it from the data using str.replace.
And print the first ten rows of movie titles again. 
"""

for i in range(df.shape[0]):
    df.movie_title[i] = df.movie_title[i].replace("Ê", "")
```

```python
df
```

Output:
```
movie_title  director_id  country_id content_rating  \
0    The Shawshank Redemption           34           1              R   
1              The Green Mile           34           1              R   
2               The Godfather           33           1              R   
3      The Godfather: Part II           33           1              R   
4              Apocalypse Now           33           1              R   
..                        ...          ...         ...            ...   
173        Dancer in the Dark           66          13              R   
174           Pan's Labyrinth           42          17              R   
175            The Sea Inside            2          17          PG-13   
176  The Secret in Their Eyes           61          16              R   
177        Mad Max: Fury Road           39          19              R   

     title_year  imdb_score      gross  duration  id_x    country  id_y  \
0          1994         9.3   28341469       142     1        USA 
```

<!-- grade_id: q3_test, points: 1, grade: true, solution: false -->
```python
assert_frame_equal(df, sol.df)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q4, points: 0, grade: false, solution: true -->
```python
""" Q4:
Who is the director with the most movies? First get the number of movies per "director_name", then save the director's name
and count as a series of length 1 called "director_with_most"
"""


df.pivot_table(index="director_name", aggfunc="size")
df.director_name.value_counts()
pd.Series(df.groupby(["director_name"], sort = True).size())
```

Output:
```
director_name
Akira Kurosawa           1
Alejandro Amenabar       1
Alejandro G. Inarritu    2
Alfred Hitchcock         1
Andrew Stanton           2
                        ..
Victor Fleming           2
Vincent Paronnaud        1
Wes Anderson             1
Wolfgang Petersen        1
Woody Allen              1
Length: 117, dtype: int64
```

```python
df.director_name.value_counts()
```

Output:
```
Christopher Nolan      7
Martin Scorsese        6
Quentin Tarantino      6
Steven Spielberg       5
Ridley Scott           4
                      ..
Morten Tyldum          1
Fritz Lang             1
John G. Avildsen       1
James McTeigue         1
Oliver Hirschbiegel    1
Name: director_name, Length: 117, dtype: int64
```

```python
director_with_most = pd.Series(df.director_name.value_counts()[0], index=[df.director_name.value_counts().index[0]], name="director_name")
```

```python
sol.director_with_most
```

Output:
```
Christopher Nolan    7
Name: director_name, dtype: int64
```

```python

```

<!-- grade_id: q4_test, points: 1, grade: true, solution: false -->
```python
assert_series_equal(director_with_most, sol.director_with_most)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q5, points: 0, grade: false, solution: true -->
```python
"""Q5:
Save all of this director's movies and their ratings in a variable called all_movies_ratings, then print this variable.
(The director with the most movies you got from the last question.)
"""

all_movies_ratings = df[df.director_name == "Christopher Nolan"].loc[:,["movie_title", "imdb_score"]]


```

```python
sol.all_movies_ratings
```

Output:
```
movie_title  imdb_score
5         The Dark Knight         9.0
6               Inception         8.8
7            Interstellar         8.6
8                 Memento         8.5
9            The Prestige         8.5
10  The Dark Knight Rises         8.5
11          Batman Begins         8.3
```

<!-- grade_id: q5_test, points: 1, grade: true, solution: false -->
```python
assert_frame_equal(all_movies_ratings, sol.all_movies_ratings)
print("Success!")
```

Output:
```
Success!
```

```python
pd.DataFrame(pd.DataFrame(df.iloc[rint, :]))
```

Output:
```
54    Alien
Name: movie_title, dtype: object
```

```python
rand_goodmovie
```

Output:
```
158
movie_title        The Pianist
director_id                 94
country_id                   7
content_rating               R
title_year                2002
imdb_score                 8.5
gross                 32519322
duration                   150
id_x                         7
country                 France
id_y                        94
director_name   Roman Polanski
```

```python
random_title = rand_goodmovie.iloc[0]
random_imdb_score = rand_goodmovie.iloc[5]
```

```python
random_title.iloc[0]
```

Output:
```
'The Pianist'
```

<!-- grade_id: q6, points: 0, grade: false, solution: true -->
```python
"""Q6:
Recommend a **random** movie that has a rating of over 8.3. 
Store the random recommendation in a variable called "rand_goodmovie".
What is the title and imdb_score of your recommendation?
 
Here are the steps:
1. Query the data ('df' DataFrame) for movies with a rating over 8.3 (imdb_score > 8.3)
2. Create a random integer index location to get a single movie recommendation
3. Save the random movie recommendation in a DataFrame called 'rand_goodmovie'
4. Save the title of the random movie recommendation in a variable "random_title" and print it
5. Save the imdb_score of the random movie recommendation in a variable "random_imdb_score" and print it

"""
# Do not modify this part, it's needed for grading
import random
random.seed(0)
df = df[df.imdb_score > 8.3]
rint = random.randint(0, df.shape[0])


rand_goodmovie = pd.DataFrame(pd.DataFrame(df.iloc[rint, :])).T
random_title = rand_goodmovie.iloc[:, 0]
random_imdb_score = rand_goodmovie.iloc[:, 5]
```

```python
rand_goodmovie
```

Output:
```
movie_title director_id country_id content_rating title_year imdb_score  \
158  The Pianist          94          7              R       2002        8.5   

        gross duration id_x country id_y   director_name  
158  32519322      150    7  France   94  Roman Polanski
```

```python
random_title
```

Output:
```
158    The Pianist
Name: movie_title, dtype: object
```

```python
random_imdb_score
```

Output:
```
158    8.5
Name: imdb_score, dtype: object
```

<!-- grade_id: q6_test, points: 3, grade: true, solution: false -->
```python
from nose.tools import assert_in
assert_in(rand_goodmovie[["movie_title", "imdb_score"]].values, sol.possible_goodmovies[["movie_title", "imdb_score"]].values)
assert_equal(random_title.iloc[0], rand_goodmovie["movie_title"].iloc[0])
assert_equal(random_imdb_score.iloc[0], rand_goodmovie["imdb_score"].iloc[0])
print("Success!")

```

Output:
```
Success!
```

```python

```